```python
from __future__ import print_function

%matplotlib inline
%precision 8

import numpy
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

**My Name: Vincent Zhao**

# HW 2: BVP Problems 1

## Question 1

We can derive a non-uniform 3-point approximation to $u''(x_i)$ with the stencil
$$
    u''(x_i) \approx \frac{2}{\Delta x_i ( \Delta x_i + \Delta x_{i+1} )} U_{i-1} - \frac{2}{\Delta x_i \Delta x_{i+1}} U_{i} + \frac{2}{\Delta x_{i+1} ( \Delta x_i + \Delta x_{i+1} )} U_{i+1}
$$
where $\Delta x_i = x_i - x_{i-1}$.

**(a)** [15] Show that the truncation error for this approximation is 
$$
    \tau_i = \frac{1}{3} (\Delta x_{i+1} - \Delta x_i) u^{(3)}(x_i) + \frac{1}{12} \left (\frac{\Delta x_i^3 + \Delta x_{i+1}^3}{\Delta x_i + \Delta x_{i+1}} \right ) u^{(4)}(x_i) + \cdots.
$$
This leads us to believe that using this approach would be only first order accurate as 
$$
    \tau_i = \mathcal{O}(\max (\Delta x_i, \Delta x_{i+1} ) )
$$
if $\Delta x_i \neq \Delta x_{i+1}$.

<!-- grade_id: cell-4030827056886998, points: 15, grade: true, solution: true -->
**Answer to Question 1 (a)**

Since:
$$
    u''(x_i) \approx \frac{2}{\Delta x_i ( \Delta x_i + \Delta x_{i+1} )} U_{i-1} - \frac{2}{\Delta x_i \Delta x_{i+1}} U_{i} + \frac{2}{\Delta x_{i+1} ( \Delta x_i + \Delta x_{i+1} )} U_{i+1}
$$

We can know $\tau_i$ should be goes like the equation below:

$$
    \tau_i = \frac{2}{\Delta x_i ( \Delta x_i + \Delta x_{i+1} )} U_{i-1} - \frac{2}{\Delta x_i \Delta x_{i+1}} U_{i} + \frac{2}{\Delta x_{i+1} ( \Delta x_i + \Delta x_{i+1} )} U_{i+1} - f(x_i)
$$
Where $f(x) = u^{\prime\prime}(x)$

We also assume $u(x)$ is sommth and we can replace the $u(x_i)$ wieh their Taylor series counterparts below:
$$
    u(x_{i \pm 1}) = u(x_i) \pm u'(x_i) \Delta x + \frac{1}{2} u''(x_i) \Delta x^2 \pm \frac{1}{6} u'''(x_i) \Delta x^3 + \frac{1}{24} u^{(4)}(x_i) \Delta x^4 + \mathcal{O}(\Delta x^5)
$$

We then can have the expression of $u(x_{i+1})$ and $u(x_{i-1})$ separately:
$$
    u(x_{i + 1}) = u(x_i) + u'(x_i) \Delta x + \frac{1}{2} u''(x_i) \Delta x^2 + \frac{1}{6} u'''(x_i) \Delta x^3 + \frac{1}{24} u^{(4)}(x_i) \Delta x^4 + \mathcal{O}(\Delta x^5)
$$

$$
    u(x_{i - 1}) = u(x_i) - u'(x_i) \Delta x + \frac{1}{2} u''(x_i) \Delta x^2 - \frac{1}{6} u'''(x_i) \Delta x^3 + \frac{1}{24} u^{(4)}(x_i) \Delta x^4 + \mathcal{O}(\Delta x^5)
$$


We also need to substitute these given info $\Delta x_i \neq \Delta x_{i+1}$ and $\Delta x_i = x_i - x_{i-1}$ into our Taylor series and get:
$$
    u(x_{i + 1}) = u(x_i) + u'(x_i) \Delta x_{i+1} + \frac{1}{2} u''(x_i) \Delta x^2_{i+1} + \frac{1}{6} u'''(x_i) \Delta x^3_{i+1} + \frac{1}{24} u^{(4)}(x_i) \Delta x^4_{i+1} + \mathcal{O}(\Delta x^5_{i+1})
$$

$$
    u(x_{i - 1}) = u(x_i) - u'(x_i) \Delta x_{i} + \frac{1}{2} u''(x_i) \Delta x^2_{i} - \frac{1}{6} u'''(x_i) \Delta x^3_{i} + \frac{1}{24} u^{(4)}(x_i) \Delta x^4_{i} + \mathcal{O}(\Delta x^5_{i})
$$

Then, $\tau$ can be rewritten as:
$$
\begin{aligned}
    \tau_i &= \frac{2}{\Delta x_i ( \Delta x_i + \Delta x_{i+1} )} u(x_{i-1}) - \frac{2}{\Delta x_i \Delta x_{i+1}} u(x_{i}) + \frac{2}{\Delta x_{i+1} ( \Delta x_i + \Delta x_{i+1} )} u(x_{i+1}) - f(x_i)\\
    &= u^{\prime\prime} + \frac{u^{(3)}(x_i)\Delta x^3_{i+1}}{3\Delta x_{i+1} (\Delta x_i + \Delta x_{i+1}) } 
    - \frac{u^{(3)}(x_i)\Delta x^3_{i}}{3\Delta x_{i} (\Delta x_i + \Delta x_{i+1}) } 
    + \frac{u^{(4)}(x_i)\Delta x^4_{i+1}}{12\Delta x_{i+1} (\Delta x_i + \Delta x_{i+1}) }
    + \frac{u^{(4)}(x_i)\Delta x^4_{i}}{12\Delta x_{i} (\Delta x_i + \Delta x_{i+1}) }
    - f(x_i) + ...\\
    &= \frac{1}{3} (\Delta x_{i+1} - \Delta x_i) u^{(3)}(x_i) + \frac{1}{12} \left (\frac{\Delta x_i^3 + \Delta x_{i+1}^3}{\Delta x_i + \Delta x_{i+1}} \right ) u^{(4)}(x_i) + \cdots.
\end{aligned}
$$

**(b)** [20] Implement a finite difference method based on the above discretization for the problem
$$
    u''(x) = e^x ~~~~ \Omega = [0, 1] \\
    u(0) = \alpha = 0 ~~~~ u(1) = \beta = 3
$$
whose true solution is
$$
    u(x) = e^x + (\sigma + 1) \times (x - 3) + \beta - e^3.
$$
Write the function so that it can take in a parameter `grid_type` that determines the type of grid to use.

```python
def grid_points(a, b, m, grid_type='uniform'):
    """Returns a set of m grid points between a and b
    
    The parameter `grid_type` determines what the distribution of 
    points is.
     - grid_type = 'uniform' - Uniform distribution of points
     - grid_type = 'right_layer' - Clustered near right boundary
     - grid_type = 'random' - Randomly distributed
     - grid_type = 'chebyshev' - Chebyshev extreme points
    """
    # Computational grid
    z = numpy.linspace(0, 1, m + 2)
    
    if grid_type.lower() == "uniform":
        x = a + (b - a) * z
    elif grid_type.lower() == "right_layer":
        x = a + (b - a) * (1.0 - (1.0 - z)**2)
    elif grid_type.lower() == "random":
        x = numpy.empty(m + 2)
        x[1:-1] = a + (b - a) * numpy.sort(numpy.random.uniform(size=(m)))
        x[0] = a
        x[-1] = b
    elif grid_type.lower() == "chebyshev":
        x = a + (b - a) * 0.5 * (1.0 + numpy.cos(numpy.pi * (1.0 - z)))
    else:
        raise ValueError("Unknown grid type.")
    
    return x
```

<!-- grade_id: cell-5057675592880702, points: 0, grade: false, solution: true -->
```python
def solve_nonuniform(a, b, alpha, beta, m, f, grid_type='uniform'):
    """Solve the ODE u'' = f(x)
    
    This function solves the ODE above using the grid specfied on the
    interval x \in [a, b] with boundary conditions u(a) = alpha and
    u(b) = beta.
    
    :Input:
     - *a* (float) - The beginning of the interval x \in [a, b]
     - *b* (float) - The end of the interval x \in [a, b]
     - *alpha* (float) - Value that defines the BC at x = a
     - *beta* (float) - Value that defines the BC at x = b
     - *m* (int) - Number of grid points to use
     - *f* (func) - Right hand side function defined in the ODE
     - *grid_type* (string) - Type of grid to use
     
    :Output:
     - *x* (ndarray) - Location of grid points used including boundaries
     - *U* (ndarray) - Solution array including boundary conditions (m + 2)
    """
    
    x = grid_points(a, b, m, grid_type) 
    
    # construct matrix A
    A = numpy.zeros((m + 2, m + 2))
    
    for i in range(1, m + 1):
        delta_i = x[i] - x[i-1]
        delta_i1 = x[i+1] - x[i]
        A[i,i] = -2.0 / (delta_i * delta_i1)
        A[i,i-1] = 2.0 / (delta_i * (delta_i + delta_i1))
        A[i,i+1] = 2.0 / (delta_i1 * (delta_i + delta_i1))
    
    # Boundary conditions
    A[0, 0] = 1.0
    A[0, 1] = 0.0
    A[-1, -1] = 1.0
    A[-1, -2] = 0.0
    
    # Construct RHS
    b = numpy.zeros(m+2)
    x1 = x[1:-1]
    b[1:-1] = f(x1)
    b[0] = alpha 
    b[-1] = beta 
    
    # Solve system
    U = numpy.linalg.solve(A, b)
    
    
#     # YOUR CODE HERE
#     raise NotImplementedError()
    return x, U
```

<!-- grade_id: cell-1271126761078309, points: 20, grade: true, solution: false -->
```python
a = 0.0
b = 1.0
alpha = 0.0
beta = 3.0
f = lambda x: numpy.exp(x)
u_true = lambda x: (4.0 - numpy.exp(1.0)) * x - 1.0 + numpy.exp(x)
x, U = solve_nonuniform(a, b, alpha, beta, 100, f, grid_type='uniform')
print(numpy.linalg.norm(u_true(x) - U, ord=2))
fig = plt.figure()
axes = fig.add_subplot(1, 1, 1)
axes.plot(x, U, 'bo')
axes.plot(x, u_true(x), 'k')
axes.set_xlabel("x")
axes.set_ylabel("u(x)")
axes.set_xlim((a, b))
assert(numpy.linalg.norm(u_true(x) - U, ord=2) < 1e-4)
print("Success!")
```

**(c)** [10] We would expect in general then that we would see only 1st order convergence if we let `grid_type = 'random'`.  Analyze the convergence of the method using each of the types of point distributions available in the function `grid_points`.  Use the $\Delta x$ computed from the uniform grid point distribution.

<!-- grade_id: cell-8633857151302828, points: 10, grade: true, solution: true -->
```python
# Problem setup
a = 0.0
b = 1.0
alpha = 0.0
beta = 3.0
f = lambda x: numpy.exp(x)
u_true = lambda x: (4.0 - numpy.exp(1.0)) * x - 1.0 + numpy.exp(x)

#compute the error as a function of delta_X
m_range = numpy.arange(10, 1000, 10)
delta_x = numpy.empty(m_range.shape)
error = numpy.empty((m_range.shape[0],4))

grid_type = ["uniform", "right_layer", "random", "chebyshev"]

for (i, m) in enumerate(m_range):
    
    delta_x[i] = (b - a) / (m + 1)
    
    #compute solution
    for k in range(4):
        x, U = solve_nonuniform(a, b, alpha, beta, m, f, grid_type=grid_type[k])
        error[i, k] = numpy.linalg.norm(U - u_true(x), ord=2) / numpy.linalg.norm(u_true(x),ord =2)

titles = ["uniform", "right_layer", "random", "chebyshev"]
order_C = lambda delta_x, error, order: numpy.exp(numpy.log(error) - order * numpy.log(delta_x))

for i in range(4):
    fig = plt.figure()
    axes = fig.add_subplot(1, 1, 1)

    axes.loglog(delta_x, error[:, i], 'ko', label="Approx. Derivative")

    axes.loglog(delta_x, order_C(delta_x[0], error[0,i], 1.0) * delta_x**1.0, 'r--', label="1st Order")
    axes.loglog(delta_x, order_C(delta_x[0], error[0,i], 2.0) * delta_x**2.0, 'b--', label="2nd Order")
    
    axes.legend(loc=4)
    axes.set_title(titles[i])
    axes.set_xlabel("$\Delta x$")
    axes.set_ylabel("$|u(x) - U|$")

plt.show()


# # YOUR CODE HERE
# raise NotImplementedError()
```

## Question 2

Consider the convection-diffusion problem defined by
$$
    u_t + au_x = \kappa u_{xx} + \psi(x)
$$
This is similar to the heat conduction problem except that we also have a convective term $a~u_x$.  You can think of $u(x,t)$ as describing the distribution of dye in a stream flowing at velocity $a$ where $\kappa$ describes how quickly the dye diffuses and $\psi(x)$ represents dye being added to the stream.

Considering the steady state solution of the problem above leads to the ODE
$$
    \epsilon u''(x) - u'(x) = f(x)
$$
where $\epsilon$ represents the relative strength of the advective and diffusive effects.

We will solve this problem with the domain and boundary conditions
$$
    \Omega = [0, 1] ~~~~ u(0)=\alpha ~~~~ u(1)=\beta ~~~~ f(x) = -1
$$
whose true solution is
$$
    u(x) = \alpha + x + (\beta - \alpha -1) \left( \frac{e^{x / \epsilon} - 1}{e^{1 / \epsilon} - 1} \right ).
$$

**(a)** [10] Discretize the boundary value problem presented using at least second order finite difference methods.  Make sure to pick the appropriate finite difference approach at the boundaries.

<!-- grade_id: cell-484106738343592, points: 0, grade: false, solution: true -->
```python
def solve_adv_diff(alpha, beta, epsilon, m):
    """Solve the steady-state advection-diffusion problem
    
    :Input
     - *alpha* (float) - Value that defines the BC at x = 0
     - *beta* (float) - Value that defines the BC at x = 1
     - *epsilon* (float) - The ratio of the advection speed and diffusivity
     - *m* (int) - Number of grid points to use
     
    :Output:
     - *x* (ndarray) - Location of grid points used including boundaries
     - *U* (ndarray) - Solution array including boundary conditions (m + 2)
    """
    
    # Descretization
    x = numpy.linspace(0., 1., m + 2)
    delta_x = 1.0 / (m + 1)
    f = lambda x: -1.

    # Matrix Construction
    A = numpy.zeros((m, m))
    diagonal = epsilon * numpy.ones(m) / delta_x**2
    A += numpy.diag(diagonal * -2.0, 0)
    A += numpy.diag(diagonal[:-1] - 1.0/(2.0 * delta_x), 1)
    A += numpy.diag(diagonal[:-1] + 1.0/(2.0 * delta_x), -1)

    # Construct RHS
    b = f(x) * numpy.ones(m)
    b[0] -= epsilon * alpha/delta_x**2 + alpha/(2 * delta_x)
    b[-1] -= epsilon * beta/delta_x**2 - beta/(2 * delta_x)
    
    # Solve system
    U = numpy.empty(m + 2)
    U[0] = alpha
    U[-1] = beta
    U[1:-1] = numpy.linalg.solve(A, b)
    
#     # YOUR CODE HERE
#     raise NotImplementedError()
    return x, U
```

<!-- grade_id: cell-6991656051774681, points: 10, grade: true, solution: false -->
```python
alpha = 1.0
beta = 3.0
epsilon = -0.1
u_true = lambda x: alpha + x + (beta - alpha - 1.0) * (numpy.exp(x / epsilon) - 1.0) / (numpy.exp(1.0 / epsilon) - 1.0)
x, U = solve_adv_diff(alpha, beta, epsilon, 100)

fig = plt.figure()
axes = fig.add_subplot(1, 1, 1)
axes.plot(x, U, 'bo')
axes.plot(x, u_true(x), 'k')
axes.set_xlabel("x")
axes.set_ylabel("u(x)")
axes.set_xlim((0.0, 1.0))

error = numpy.linalg.norm((x[1] - x[0]) * (U - u_true(x)), ord=1)
print(error)
assert(error < 1e-4)
print("Success!")
```

**(b)** [10] Using your function from part (a) explore the parameter space a bit by varying $m$ and $\epsilon$.  Suggested values to explore include $m \in [10, 100]$ and $\epsilon \in [-0.01, -1.0]$.

Describe what you observe with both plots and written observations.  In particular do your solutions seem physically reasonable?  If your code has difficulties explain why it might be having issues.

<!-- grade_id: cell-4718808055966480, points: 5, grade: true, solution: true -->
```python
a = 0.0
b = 1.0
alpha = 1.0
beta = 3.0
epsilon = -0.1
u_true = lambda x: alpha + x + (beta - alpha - 1.0) * (numpy.exp(x / epsilon) - 1.0) / (numpy.exp(1.0 / epsilon) - 1.0)

# # YOUR CODE HERE
# raise NotImplementedError()

x, U = solve_adv_diff(alpha, beta, epsilon=-0.01, m=10)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.01, m = 10$')
plt.xlabel("$x$")
plt.ylabel("$u(x)$",rotation = 0)
plt.legend(loc = 'best')
plt.grid()
plt.show()

x, U = solve_adv_diff(alpha, beta, epsilon=-0.01, m=50)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.01, m = 50$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()


x, U = solve_adv_diff(alpha, beta, epsilon=-0.01, m=100)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.01, m = 100$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()


x, U = solve_adv_diff(alpha, beta, epsilon=-0.05, m=10)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.1, m = 10$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()

x, U = solve_adv_diff(alpha, beta, epsilon=-0.05, m=50)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.1, m = 50$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()

x, U = solve_adv_diff(alpha, beta, epsilon=-0.05, m=100)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.1, m = 100$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()

x, U = solve_adv_diff(alpha, beta, epsilon=-0.1, m=10)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.1, m = 10$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()

x, U = solve_adv_diff(alpha, beta, epsilon=-0.1, m=50)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.1, m = 50$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()

x, U = solve_adv_diff(alpha, beta, epsilon=-0.1, m=100)
plt.plot(x, U, 'b', alpha=0.8, label='Approximate Solution')
plt.plot(x, u_true(x), 'k', alpha=1, label='True Solution')
plt.title(r'$\epsilon = -0.1, m = 100$')
plt.xlabel("x")
plt.ylabel("u(x)")
plt.legend()
plt.grid()
plt.show()

```

<!-- grade_id: cell-8000155034097309, points: 5, grade: true, solution: true -->
**Answer to Question 2**

I tried three $\epsilon$ and three $m$ respectively and found something below:
- $\epsilon$ controls whether our approximate solution is close to the true solution, when $\epsilon$ goes from $-0.01$ to $-0.1$, my approximate solution goes closer to the true solution
- $m$ controls whether the approximate solution is smooth enough, which means how many jumps or steps or points we need to use to simulate the true solution. When m goes larger, the curve of approximate solution will be smoother

## Question 3

Consider the following BVP:
$$
    u''(x) + u(x) = 0 \quad \quad x \in \Omega = (a, b) \\
    u(a) = \alpha \quad u(b) = \beta
$$
where $L = b - a$.

**(a)** [10] What is the exact solution of the BVP?

<!-- grade_id: cell-3539711075833027, points: 10, grade: true, solution: true -->
We know from the equation $u''(x) + u(x) = 0$ that the exact solution of the BVP could be:
$$u(x) = c_1 cos(x) + c_2 sin(x)$$

Then, we can derive $c_1$ and $c_2$ from the given boundary condition below:
$$u(a) = \alpha \quad u(b) = \beta$$

First, substitute (a,$\alpha$) and (b,$\beta$) into our solution, we can get:
$$
 \alpha = u(a) = c_1 cos(a) + c_2 sin(a) \quad \quad (1)\\
 \beta = u(b) = c_1 cos(b) + c_2 sin(b) \quad \quad (2)
$$

$(1) \times sin(b) - (2) \times sin(a)$ we can get $c_1$:

$$
c_1 = \frac{\beta sin(a) - \alpha sin(b)}{cos(b)sin(a) - cos(a)sin(b)}
$$


$(1) \times cos(b) - (2) \times cos(a)$ we can then get $c_2$:

$$
c_2 = \frac{\beta cos(a) - \alpha cos(b)}{sin(b)cos(a) - sin(a)cos(b)}
$$

$u(x)$ can then be rewritten as:
$$
u(x) = \frac{\beta sin(a) - \alpha sin(b)}{cos(b)sin(a) - cos(a)sin(b)} cos(x) + \frac{\beta cos(a) - \alpha cos(b)}{sin(b)cos(a) - sin(a)cos(b)} sin(x)
$$

**(b)** [15] Write a function that solves this equation using a finite difference approximation.

<!-- grade_id: cell-928726382559822, points: 0, grade: false, solution: true -->
```python
def solve_linear_pendulum(a, b, alpha, beta, m):
    """Solve u'' + u = 0 with x \in (a, b) and u(a) = \alpha, u(b) = \beta
    
    Returns the points the solution was approximated at including the boundaries and
    the solution
    """
    # Descretization
    x = numpy.linspace(a, b, m + 2)
    delta_x = (b - a) / (m + 1)
    
    # Matrix Construction
    A = numpy.zeros((m, m))
    diagonal = numpy.ones(m)
    A += numpy.diag((delta_x**2 - 2.) * diagonal, 0)
    A += numpy.diag(diagonal[:-1], 1)
    A += numpy.diag(diagonal[:-1], -1)
    
    # Construct RHS
    b = numpy.zeros(m)
    b[0] -= alpha
    b[-1] -= beta
    
    # Solve system
    U = numpy.empty(m + 2)
    U[0] = alpha 
    U[-1] = beta 
    U[1:-1] = numpy.linalg.solve(A, b)
    
    
#     # YOUR CODE HERE
#     raise NotImplementedError()
    return x, U
```

```python
a = 0.0
b = 1.0
alpha = 2.0
beta = 3.0
C1 = (alpha - beta * numpy.sin(a) / numpy.sin(b)) / (numpy.cos(a) - numpy.sin(a) / numpy.tan(b))
C2 = beta / numpy.sin(b) - C1 / numpy.tan(b)
u_true = lambda x: C1 * numpy.cos(x) + C2 * numpy.sin(x)
x, U = solve_linear_pendulum(a, b, alpha, beta, 200)
fig = plt.figure()
axes = fig.add_subplot(1, 1, 1)
axes.plot(x, u_true(x), 'k')
axes.plot(x, U, 'ob')
axes.set_xlim(a, b)
axes.set_xlabel("x")
axes.set_ylabel("u(x)")
error = numpy.linalg.norm((x[1] - x[0]) * (U - u_true(x)), ord=1)
print("Error = ", error)
assert(error < 1e-5)
print("Success!")
```

<!-- grade_id: cell-3018382947949415, points: 15, grade: true, solution: false -->
```python
a = 0.0
b = 1.0
alpha = 2.0
beta = 3.0
C1 = (alpha - beta * numpy.sin(a) / numpy.sin(b)) / (numpy.cos(a) - numpy.sin(a) / numpy.tan(b))
C2 = beta / numpy.sin(b) - C1 / numpy.tan(b)
u_true = lambda x: C1 * numpy.cos(x) + C2 * numpy.sin(x)
x, U = solve_linear_pendulum(a, b, alpha, beta, 100)
fig = plt.figure()
axes = fig.add_subplot(1, 1, 1)
axes.plot(x, u_true(x), 'k')
axes.plot(x, U, 'ob')
axes.set_xlim(a, b)
axes.set_xlabel("x")
axes.set_ylabel("u(x)")
error = numpy.linalg.norm((x[1] - x[0]) * (U - u_true(x)), ord=1)
print("Error = ", error)
assert(error < 1e-5)
print("Success!")
```

**(c)** [5] Let $a=0$ and $b=\pi$.  For what values of $\alpha$ and $\beta$ does this BVP have solutions?  Sketch out a family of solutions in the case of infinite numbers of solutions.

<!-- grade_id: cell-5480550098808611, points: 2.5, grade: true, solution: true -->
**Answer to Question 3 (c)**

When $a = 0$ and $b = \pi$, substitute them into our $u(x)$ we then can get:

$$
u(0) = c_1 = \alpha\\
u(\pi) = -c_1 = \beta
$$

We can only have solutions when $\alpha = -\beta$

```python
def solve_linear_pendulum(a, b, alpha, beta, m):
    """Solve u'' + u = 0 with x \in (a, b) and u(a) = \alpha, u(b) = \beta
    
    Returns the points the solution was approximated at including the boundaries and
    the solution
    """
    # Descretization
    x = numpy.linspace(a, b, m + 2)
    delta_x = (b - a) / (m + 1)
    
    # Matrix Construction
    A = numpy.zeros((m, m))
    diagonal = numpy.ones(m)
    A += numpy.diag((delta_x**2 - 2.) * diagonal, 0)
    A += numpy.diag(diagonal[:-1], 1)
    A += numpy.diag(diagonal[:-1], -1)
    
    # Construct RHS
    b = numpy.zeros(m)
    b[0] -= alpha
    b[-1] -= beta
    
    # Solve system
    U = numpy.empty(m + 2)
    U[0] = alpha 
    U[-1] = beta 
    U[1:-1] = numpy.linalg.solve(A, b)
    
    
#     # YOUR CODE HERE
#     raise NotImplementedError()
    return x, U
```

<!-- grade_id: cell-7949057030139063, points: 2.5, grade: true, solution: true -->
```python
a = 0.0
b = numpy.pi

for c in range(-5,5):
    alpha = c
    beta = -c
    C1 = (alpha - beta * numpy.sin(a) / numpy.sin(b)) / (numpy.cos(a) - numpy.sin(a) / numpy.tan(b))
    C2 = beta / numpy.sin(b) - C1 / numpy.tan(b)
    u_true = lambda x: C1 * numpy.cos(x) + C2 * numpy.sin(x)
    x, U = solve_linear_pendulum(a, b, alpha, beta, 100)
    fig = plt.figure()
    axes = fig.add_subplot(1, 1, 1)
    axes.set_title("Solution and approximate solution when alpha = {}".format(c))
    axes.plot(x, u_true(x), 'k', label = 'u_true')
    axes.plot(x, U, 'ob', label = 'approximate')
    axes.legend(loc='best')
    axes.set_xlim(a, b)
    axes.set_xlabel("x")
    axes.set_ylabel("u(x)")

# # YOUR CODE HERE
# raise NotImplementedError()
```

**(d)** [5] What happens numerically for the values $a = 0$, $b = \pi$, $\alpha = 1$, $\beta = -1$ as $\Delta x \rightarrow 0$?  Also try with $\alpha = \beta = 1$.

<!-- grade_id: cell-9163036478906246, points: 2.5, grade: true, solution: true -->
**Answer to Question 3 (d)**

when $a = 0$, $b = \pi$, $\alpha = 1$, $\beta = -1$ as $\Delta x \rightarrow 0$ we can know that the exact value and approximate solution are very close to each other

<!-- grade_id: cell-5755573651360367, points: 2.5, grade: true, solution: true -->
```python
a = 0.0
b = numpy.pi

alpha = 1
beta = -1

C1 = (alpha - beta * numpy.sin(a) / numpy.sin(b)) / (numpy.cos(a) - numpy.sin(a) / numpy.tan(b))
C2 = beta / numpy.sin(b) - C1 / numpy.tan(b)

u_true = lambda x: C1 * numpy.cos(x) + C2 * numpy.sin(x)
x, U = solve_linear_pendulum(a, b, alpha, beta, 10000)

fig = plt.figure()
axes = fig.add_subplot(1, 1, 1)
axes.set_title("Solution and approximate solution when alpha = {}, beta = {}".format(alpha,beta))
axes.plot(x, u_true(x), 'k', label = 'True Solution')
axes.plot(x, U, 'ob', label = 'approximate')
axes.legend(loc='best')
axes.set_xlim(a, b)
axes.set_xlabel("x")
axes.set_ylabel("u(x)")

error = numpy.linalg.norm((x[1] - x[0]) * (U - u_true(x)), ord=1)
print("Error = ", error)

# # YOUR CODE HERE
# raise NotImplementedError()
```

```python
a = 0.0
b = numpy.pi

alpha = 1
beta = 1

C1 = (alpha - beta * numpy.sin(a) / numpy.sin(b)) / (numpy.cos(a) - numpy.sin(a) / numpy.tan(b))
C2 = beta / numpy.sin(b) - C1 / numpy.tan(b)

u_true = lambda x: C1 * numpy.cos(x) + C2 * numpy.sin(x)
x, U = solve_linear_pendulum(a, b, alpha, beta, 10000)

fig = plt.figure()
axes = fig.add_subplot(1, 1, 1)
axes.set_title("Solution and approximate solution when alpha = {}, beta = {}".format(alpha,beta))
axes.plot(x, u_true(x), 'k', label = 'True Solution')
axes.plot(x, U, 'ob', label = 'approximate')
axes.legend(loc='best')
axes.set_xlim(a, b)
axes.set_xlabel("x")
axes.set_ylabel("u(x)")

error = numpy.linalg.norm((x[1] - x[0]) * (U - u_true(x)), ord=1)
print("Error = ", error)

# # YOUR CODE HERE
# raise NotImplementedError()
```

**Answer to Question 3 (d)**

when $a = 0$, $b = \pi$, $\alpha = 1$, $\beta = 1$ as $\Delta x \rightarrow 0$ we can easily see the difference between approcimate value and the "true solution"(its not true of course! lol)